```python
# Initialize Otter
import otter
grader = otter.Notebook("HW0.ipynb")
```

# EECS C106A Python Bootcamp
### EECS C106A: Introduction to Robotics, Fall 2026

# Introduction

We will be using the Python programming language for lab and homework assignments in EECS/BioE/MechE C106A/206A. Some homework assignments will entail matrix calculations where Python will come in handy, but you are welcome to use something like Matlab instead. This notebook is meant to be a mini bootcamp on Python for students who have had experience programming in another language already (e.g. Matlab) or need a quick refresher on Python. We will be using a few popular libraries (numpy, scipy) that are very useful. If you have experience with Matlab but not Python, we recommend checking out the [numpy for Matlab users guide](https://docs.scipy.org/doc/numpy/user/numpy-for-matlab-users.html).

# Table of Contents
* [IPython Basics](#IPython-Basics)
* [Python Data Types](#Data-Types)
* [Python Functions](#Functions)

# IPython Basics

For those who have never used IPython, it is a command shell for interactive Python programming. You can imagine it as allowing you to run blocks of Python code that you would execute in a single python script (python [script_name].py) in the terminal. Benefits of using IPython include easier visualization and debugging. The purpose of this bootcamp in IPython is to give you an idea of basic Python syntax and semantics. For labs you are still expected to write and execute actual Python scripts to work with ROS.

### Executing Cells

ipython notebooks are constituted of cells, that each have text, code, or html scripts. To run a cell, click on the cell and press Shift+Enter or the run button on the toolbar. Run the cell below, you should expect an output of 6:

```python
1+2+3
```

### Import a library

To import a certain library `x`, just type `import x`
Calling function `y` from that library is simply `x.y`
To give the library a different name (e.g. abbreviation), type `import x as z`

```python
!pip3 install numpy
import numpy as np
np.add(3, 4)

```

# Python

## Data Types

### Integers and Floats

In Python2, integer division returns the floor. In Python3, there is no floor unless you specify using double slahes. The differences between Python2 and Python3 you can [check out](https://wiki.python.org/moin/Python2orPython3).

```python
5 / 4
```

```python
5.0 / 4
```

### Booleans

Python implements all usual operators for Boolean logic. English, though, is used rather than symbols (no &, ||, etc.). Pay attention to the following syntax, try to guess what the output for each print statement should be before running the cell.

```python
print(0 == False)

t = True
print(1 == t)

print(0 != t)

print(t is not 1) 

if t is True:
    print(0 != 0)
```

### Strings

Strings are supported very well. To concatenate strings we can do the following:

```python
hello = 'hello'
robot = 'robot'

print(hello + ' ' + robot + str(1))
```

To find the length of a string use `len(...)`

```python
print(len(hello + robot))
```

## Functions

Python functions are defined using the `def` keyword. For example:

```python
def hello_robot(robot_name, yell=True):
    if yell:
        print('HELLO, %s!' % robot_name.upper())
    else:
        print('hello, %s' % robot_name)
```

```python
hello_robot('Baxter') # Prints "HELLO, BAXTER!"
hello_robot('Sawyer', yell=False)  # Prints "hello, Sawyer"
```

### Lists

A list is a mutable array of data, meaning we can alter it after insantiating it. To create a list, use the square brackets [] and fill it with elements.

Key operations:
- `'+'` appends lists
- `len(y)` to get length of list y
- `y[0]` to index into 1st element of y **Python indices start from 0
- `y[1:6]` to slice elements 1 through 5 of y

```python
y = ["Robots are c"] + [0, 0, 1]
y
```

```python
len(y)
```

```python
y[0]
```

```python
def element_slicing(lst):
    # TODO: slice the first three elements of the input list 'lst' 
    # and store in a new list called 'lst_sliced', then print the 
    # 2nd element of this new list  
    
    ...

element_slicing(y) # This should print '0'

```

```python
grader.check("element_slicing")
```

### Loops

You can loop over the elements of a list like this:

```python
robots = ['baxter', 'sawyer', 'turtlebot']
for robot in robots:
    print(robot)
# Prints "baxter", "sawyer", "turtlebot", each on its own line.
```

If you want access to the index of each element within the body of a loop, use the built-in [`enumerate`](https://docs.python.org/2.7/library/functions.html#enumerate) function:

```python
robots = ['baxter', 'sawyer', 'turtlebot']

def display_enumerate(lst):
    # TODO: Using a for loop and the python built-in enumerate function,
    # print the enumerated representation of the list
    # e.g. "#1: baxter", "#2: sawyer", "#3: turtlebot", 
    # each on its own line
    ...

display_enumerate(robots)

```

```python
grader.check("display_enumerate")
```

### Numpy Array

The numpy array is like a list with multidimensional support and more functions (which can all be found [here](https://docs.scipy.org/doc/numpy/reference/index.html)).

NumPy arrays can be manipulated with all sorts of arithmetic operations. You can think of them as more powerful lists. Many of the functions that already work with lists extend to numpy arrays.

To use functions in NumPy, we have to import NumPy to our workspace. by declaring `import numpy`, which we have done previously above in this notebook already. We typically rename `numpy` as `np` for ease of use.

### Making a Numpy Array

```python
x = np.array([[1, 2, 3], [4 , 5, 6], [7, 8, 9]])
print(x)
# x is a 3x3 matrix here
```

### Finding the shape of a Numpy Array

```python
x.shape # returns the dimensions of the numpy array
```

### Building arrays by filling in entries

Instead of writing out every entry by hand, its sometimes easier to start from an array of all zeroes or the identity matrix, then fill in the only entries you care about.

```python
# Start with a 3x3 matrix of all zeros
rot = np.zeros((3, 3))
theta = np.pi / 4

# Fill in just the entries we need
rot[0, 0] = np.cos(theta)
rot[0, 1] = -np.sin(theta)
rot[1, 0] = np.sin(theta)
rot[1, 1] = np.cos(theta)
rot[2, 2] = 1

print(rot)

# np.eye(n) gives you an n x n identity matrix to start from
# instead of zeros, which is handy when most of the diagonal
# should stay 1
g = np.eye(4)
g[0, 3] = 5   # set the x-position
g[1, 3] = 2   # set the y-position
print(g)
```

### Elementwise operations

Arithmetic operations on numpy arrays correspond to elementwise operations.

```python
print(x)
print
print(x * 5) # numpy carries operation on all elements!
```

### Matrix multiplication

```python
print(np.dot(x, x))
```

### Slicing numpy arrays

Numpy uses pass-by-reference semantics so it creates views into the existing array, without implicit copying. This is particularly helpful with very large arrays because copying can be slow. Although be wary that you may be mutating an array when you don't intend to, so make sure to make a copy in these situations.

```python
orig = np.array([0, 1, 2, 3, 4, 5])
print(orig)
```

Slicing an array is just like slicing a list

```python
sliced = orig[1:4]
print(sliced)
```

Note, since slicing does not copy the array, mutating `sliced` mutates `orig`. Notice how the 4th element in `orig` changes to 9 as well.

```python
sliced[2] = 9
print(orig)
print(sliced)
```

We should use `np.copy()` to actually copy `orig` if we don't want to mutate it. 

```python
orig = np.array([0, 1, 2, 3, 4, 5])
copy = np.copy(orig)
sliced_copy = copy[1:4]
sliced_copy[2] = 9
print(orig)
print(sliced_copy)
```

### Slicing in more than one dimension

Everything above works in 1D arrays, but numpy slicing works in multiple dimensions: arr[row_slice, col_slice]. You can also assign directly into one of these slices to set a whole block of an array all at once. This comes in handy when building bigger matrices out of smaller matrices.

```python
g = np.eye(4)
rot = np.array([[0, -1, 0], [1, 0, 0], [0, 0, 1]])
position = np.array([1, 2, 3])

# Assign the top-left 3x3 block to be our matrix rot
g[:3, :3] = rot

# Assign the top 3 entries of the last column to be our array position
g[:3, 3] = position

print(g)

# You can read a block back out the same way
print(g[:3, :3]) 
print(g[:3, 3])
```

### Reshaping arrays

.reshape() lets you view the same data but with a different shape, which comes in handy when a function expects a specific shape for its input.

```python
v = np.array([1, 2, 3])
print(v.shape)          # (3,)

v_col = v.reshape(3, 1)
print(v_col.shape)      # (3, 1)
print(v_col)

# reshape doesn't copy data -- it's a view, just like slicing
v_col[0, 0] = 99
print(v)                # v is changed too!
```

```python
A = np.array([[5, 6, 8], [2, 4, 5], [3, 1, 10]])
B = np.array([[3, 5, 0], [3, 1, 1]])

def matrix_question():
    # TODO: multiply matrix A with matrix B padded with 1's to the 
    # same dimensions as A; sum this result with the identiy matrix 
    # (you may find np.concatenate, np.vstack, np.hstack, or np.eye useful). 
    # Make sure you don't alter the original contents of B. Print the result
    ...

matrix_question()

```

```python
grader.check("matrix")
```

### Handy Numpy function: arange

We use `arange` to instantiate integer sequences in numpy arrays. It's similar to the built-in range function in Python for lists. However, it returns the result as a numpy array, rather a simple list.

`arange(0,N)` instantiates an array listing every integer from 0 to N-1.

`arange(0,N,i)` instantiates an array listing every `i` th integer from 0 to N-1 .

```python
print(np.arange(-3,4)) # every integer from -3 ... 3
```

```python
def arange_question(x):
    # TODO: print every even integer from 0 up to x multiplied by 2
    # as a NumPy array
    ...

arange_question(6)
```

```python
grader.check("arange")
```

# References
- [1] EE 120 lab1
- [2] EECS 126 Lab01
- [3] EE 16a Python Bootcamp
- [4] CS 231n Python Numpy Tutorial. [Link](http://cs231n.github.io/python-numpy-tutorial/)

## Submission

1. Restart the kernel and run all cells from top to bottom.
2. Confirm that the final check reports that all available tests pass.
3. Save **hw0/HW0.ipynb**.
4. Commit and push the notebook to your private semester repository.
5. In Gradescope, submit that repository and its `main` branch using the GitHub submission option. Do not create or upload a zip file.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for arange (points: 30)
```
>>> arange_question(0)
[0]
```
```
>>> arange_question(3)
[0 4]
```


Tests for display_enumerate (points: 20)
```
>>> display_enumerate(['robot'])
#1: robot
```
```
>>> display_enumerate([10, 20, 30, 40])
#1: 10
#2: 20
#3: 30
#4: 40
```


Tests for element_slicing (points: 20)
```
>>> element_slicing([10, 20, 30])
20
```
```
>>> element_slicing(['a', 'b', 'c', 'd', 'e'])
b
```


Tests for matrix (points: 30)
```
>>> B_before_call = np.copy(B)
>>> def _run_and_check_matrix():
...     matrix_question()
...     assert bool(np.array_equal(B, B_before_call)), 'B was mutated by matrix_question()!'
>>> _run_and_check_matrix()
[[42. 39. 14.]
 [23. 20.  9.]
 [22. 26. 12.]]
```
